---
title: 'How do you process speech data?'
description: ""
---

## Introduction to SpeechRecognition Python Library

```yaml
type: VideoExercise
key: 82952ba503
xp: 50
```

`@projector_key`
f657fa6de1364fcda13b149e26867f1e

---

## Import the speech_recognition library

```yaml
type: NormalExercise
key: 081d68fe12
xp: 100
```

To save typing out `speech_recognition` every time we want to use the library, it's common practice to import it as `sr`.

This is much like Pandas is imported as `pd` and NumPy is imported as `np`.

`@instructions`
- Import the `speech_recognition` library as `sr`

`@hint`
- In Python, it's common practice to abbreviate the name of a library to it's lowercased initials or something which is relatable to the original name.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# Importing the abbreviated version of the speech_recognition library
import ____ as ____
```

`@solution`
```{python}
# Importing the abbreviated version of the speech_recognition library
import speech_recognition as sr
```

`@sct`
```{python}

```
